# New York Philharmonic

## Due Thursday, May 12 at 8 AM

In this lab, you will analyze XML data of every one of the [New York Philharmonic](http://www.nyphil.org)'s concerts between 1963 and 1973. The data resides in the file `/data/nyphil.xml`.

Note that the same program may have been used for several concerts. So the number of times a work was _programmed_ might be different from the number of times it was _performed_.

You are highly encouraged to skim the [Beautiful Soup documentation](https://www.crummy.com/software/BeautifulSoup/bs4/doc/). Unlike most documentation, it's concise and organized!

## Question 0 (5 points)

Use [Beautiful Soup](https://www.crummy.com/software/BeautifulSoup/bs4/doc/) to read the XML data into a Python object. Store the Python object called `data`. Make sure the tests below run without any errors, as this question is autograded.

<!-- grade_id: read_data_answer, points: 0, grade: false, solution: true -->
```python
from bs4 import BeautifulSoup
fp = open('/data/nyphil.xml')
data = BeautifulSoup(fp, 'xml')

```

<!-- grade_id: read_data_tests, points: 5, grade: true, solution: false -->
```python
root = list(data.children)[0]
assert(root.name == "programs")
assert(sum(1 for _ in root.children) == 1931)
```

## Question 1 (15 points)

Which works (please give composer and title) were programmed the most times over this time period? (No explanation necessary; just print out the top works, alongside the counts of how many programs they appeared on.)

<!-- grade_id: programmed_answer, points: 15, grade: true, solution: true -->
```python
import pandas as pd
work = data.findAll('work')
newworks = []
newcomposers = []
newtitle = []
for i in range(0,len(work)):
    ID = work[i]["ID"]
    if ID != "0*":
        newcomposers.append(work[i].findAll("composerName")[0].string)
        newtitle.append(work[i].findAll("workTitle")[0].string)
dfwork = pd.DataFrame()
dfwork['Title'] = newtitle
dfwork['Composer'] = newcomposers
copyworks = workinfo.groupby(['Composer', 'Title']).size().reset_index()
copyworks.columns = ['Composer','Title', 'Count']
copyworks = copyworks.sort_values('Count', ascending= False)
copyworks.head(10)
```

Output:
```
NameErrorTraceback (most recent call last)
<ipython-input-3-6a44095aba77> in <module>()
     12 dfwork['Title'] = newtitle
     13 dfwork['Composer'] = newcomposers
---> 14 copyworks = workinfo.groupby(['Composer', 'Title']).size().reset_index()
     15 copyworks.columns = ['Composer','Title', 'Count']
     16 copyworks = copyworks.sort_values('Count', ascending= False)

NameError: name 'workinfo' is not defined
```

## Question 2 (20 points)

Which works (please give composer and title) were performed the most times over this time period? (No explanation necessary; just print out the top works, alongside the counts of how many times they were performed.)

<!-- grade_id: performed_answer, points: 20, grade: true, solution: true -->
```python
programs = data.findAll('program')
newworks = []
newcomposers = []
newtitle = []
for i in range(0,len(programs)):
    times_performed = len(programs[i].find_all("concertInfo"))
    works = programs[i].find_all('work')
    for j in range(0, len(works)):
        ID = works[j]["ID"]
        if ID != "0*":
            for l in range(0,times_performed): 
                newcomposers.append(works[j].findAll("composerName")[0].string)
                newtitle.append(works[j].findAll("workTitle")[0].string)
```

```python
dfwork = pd.DataFrame()
dfwork['Title'] = newtitle
dfwork['Composer'] = newcomposers
copyworks = dfwork.groupby(['Composer', 'Title']).size().reset_index()
copyworks.columns = ['Composer','Title', 'Count']
copyworks = copyworks.sort_values('Count', ascending= False)
copyworks.head(10)
```

## Question 3a (20 points)

Make a Pandas DataFrame, where each row is a work that was programmed by the New York Philharmonic. The columns should include the composer, work title, conductor, and the date of the performance. (Hint: You may want to look at the Pandas DataFrame to remind yourself about how to use `pd.DataFrame` to create a DataFrame from a dict.)

Please print out the first few rows of your DataFrame.

still want each row to be the time it was performed just want the first date it was performed

<!-- grade_id: to_pandas_answer, points: 20, grade: true, solution: true -->
```python
programs = data.findAll('program')
newworks = []
date = []
conductor = []
newcomposers = []
newtitle = []
for i in range(0,len(programs)):
    dateperformed= programs[i].find_all('concertInfo')[0].find_all('Date')[0].string
    works = programs[i].find_all('work')
    for j in range(0, len(works)):
        ID = works[j]["ID"]
        if ID != "0*":
            newcomposers.append(works[j].findAll("composerName")[0].string)
            newtitle.append(works[j].findAll("workTitle")[0].string)
            date.append(dateperformed)
            try:
                conductor.append(works[j].findAll("conductorName")[0].string)
            except:
                conductor.append("NO CONDUCTOR")
dfwork = pd.DataFrame()
dfwork['Title'] = newtitle
dfwork['Composer'] = newcomposers
dfwork['Date'] = date
dfwork['Conductor'] = conductor
dfwork.head(10)
```

## Question 3b (10 points)

Use the DataFrame you created above to determine Leonard Bernstein's favorite composers. That is, which composers appeared on the most programs where Bernstein was conducting?

```python
bernard = dfwork[dfwork['Conductor'] == "Bernstein, Leonard"]
favcomp = bernard.groupby('Composer').size().reset_index()
favcomp.columns = ['Composer', 'Count']
favcomp.sort_values('Count', ascending=False)
```

<!-- grade_id: bernstein_response, points: 10, grade: true, solution: true -->
His top 3 favorite composers (in order) are Mahler, Beethoven, and Tchaikovsky

## Question 4 (20 points)

For each composer, calculate the number of programs that featured one (or more) of his works.  Sort the composers in descending order of the number of programs in which they appeared.

**Think:** Why can't you just call `.groupby("composer").count()` on your Pandas DataFrame from the previous question?

<!-- grade_id: composers_response, points: 20, grade: true, solution: true -->
```python
test = dfwork.groupby(['Composer', 'Date']).count().reset_index()
test2  = test.groupby(['Composer']).size().reset_index()
test2.sort_values(0, ascending=False)
```